Circadian control of a sex-specific behaviour in Drosophila

Author response:

The following is the authors’ response to the previous reviews

Joint Public Review:

(1) Problems associated with averaging: The authors intended to focus on the oviposition clock in individual females, however due to the inherent noise in the oviposition rhythm they had to resort to averaging across Lomb-Scargle periodograms generated from individual time-series. They then tested whether the averaged periodogram contains a significant frequency. However, this reduction in noise also reduces the ability to compare differences in power of the rhythm across individuals. Furthermore, this method makes it especially difficult to distinguish the contribution of subsets of the circuit on the proportion of rhythmic flies and the power of the rhythm. In this revised version the authors use two manipulations to disrupt the molecular clock, which could have different success rates based on the type and number of cells targeted. Unfortunately, the type of averaging used prevents the detection of any such effects. It is to be noted that, indeed, individual-level differences in period between the PdfDicerGal4 > perRNAi and UAS-perRNAi lines help the authors to establish that there is a significant reduction in period length when the molecular clock is abolished in PDF cells. These individual measurements are now very helpful in discerning the effect of manipulations carried out on different circadian neural subsets, some of which could have been missed if only averages were considered.

First, it is important to emphasize that we are certainly not "averaging across LombScargle periodograms". As explained in the paper (and at length in the Supplementary Material), what we do is first to detrend each individual time series, then average _all_ the resulting time series (and not only those of rhythmic individuals), and finally take the Lomb-Scargle periodogram of this average series. Nevertheless, we agree with the reviewer in that the use of averages reduces our ability of understanding what happens at the individual level. The problem is that in most cases the presence of noise has made it difficult to draw any meaningful conclusions. One fortunate exception is the one mentioned by the reviewer. Averaging, on the other hand, has allowed us to extract some useful information in those cases.

(2) Sensitivity to sample size: Averaging reduces the effect of random background noise but noise reduction is dependent upon sample size. Comparing genotypes with different sample sizes in addition to varying signal to noise ratios (which might also change with neural manipulations) makes it difficult to estimate how much of the rhythm structure is contributed by a given neuronal subset; thus, whenever possible comparisons should be made between groups that include similar number of flies. This problem is compounded when the averaged periodogram is composed of both rhythmic and weakly rhythmic individuals. For instance, in the main text the reported value of period length of pdfDicerGal4 > perRNAi is 20.74h (see also Fig 2J) but in the Supplementary figure 2S1 this is close to 22h, while the values reported for the control are largely similar (24.35h in Fig 2H versus ~24h in Fig 2S1). A difference of 3.6h between control and experimental flies is much greater than 2h. Which estimate (average versus individual) is more reliable in predicting the behavior of these flies is difficult to determine without further experiments.

In most of the experiments analyzed for this paper the number of flies for control and experimental genotypes are very similar. In the remaining ones, the number of flies for experimental genotypes is roughly twice the number of flies for control genotypes. As mentioned, noise reduction depends on sample size. This implies that, when a genotype is assessed as rhytyhmic the sample size used is evidently large enough. On the other hand, when a genotype is assessed as arrhythmic it is important to know if sample size is large enough. It is for this reason that we have used many more flies for arrhythmic genotypes vs. their control genotypes.

Regarding the period difference between the average of rhythmic individuals, and the population denoised average, notice first that they are not necessarily excactly the same thing, since our population average uses all flies, and the denoising might introduce some variations over the underlying periods (which would be undetectable without the denoising). Also, and more importantly, Fig. 2S1 shows that for the average of the individual periods the error bars are large, and thus statistically, the reported value for the population average falls within the confidence interval for the individual average.

(3) Based on the newly provided data for individual fly periodograms the reader can visually evaluate the rhythmicity associated with each genotype. Such visual inspection did not reveal any clear difference between the proportion of rhythmic individuals between experimental and parental GAL4 and/or UAS controls, except for experiments using per01 mutant animals. This is surprising since if these circuits are controlling the oviposition rhythm, perturbing them should affect most individuals in a similar way.

The problem here is that, given the amount of noise present in this behavior, it is difficult to obtain any reliable information from individual records, since, by its random nature, in a given experiment noise might be disturbing the expected behavior of individuals in very different ways. That is the reason why we have resorted to population averages.

Other comments

Disrupting the clock in the 5th sLNv and 3 Cry+ LNds (and weakly in a small subset of DN1) affected egg-laying. Although the work emphasizes the importance of the LNd, the role of the 5th sLNv's role should be discussed.

As mentioned in the paper, what the experiments show is that the 3 Cry+ LNds and 5th sLNv (usually called E cells) are candidates to be the main drivers of the oviposition rhythm, but the connectomics show that only 2 Cry+ LNds are connected to the oviposition circuit. In order to be more accurate, throughout the corresponding section (now called "The molecular clock in E neurons is necessary for rhythmic egg-laying") of the corrected manuscript we have always referred to the cells marked by the driver as E-cells. In the Discussion, we have added a line commenting that, in the connectome, the 5th sLNv is not connected to any cells of the oviposition circuit.

Minor corrections:

In subsection "Two Cry+ LNd neurons directly oviIN", there was a mistake in the use of "E1" and "E2" (their meanings were interchanged). We have corrected this section, giving the correct definitions. We have also corrected some minor english typos.

Joint Recommendations for the authors:

(1) Line 234 'to disrupt the molecular clock in (those) neurons', Please clearly describe the cell types in which MB122B driver works.

We have clarified the cell types in which MB122B driver is expressed (line 236)

(2) Line 235 gen cycle, should be gen'e' cycle

The typo has been corrected

(3) The authors should provide the raw data in repositories as per journal policy of eLife.

The data are now available at the following links:

https://github.com/srisaug/flywork/blob/main/RawData_Rivaetal_eLife2025_Fig4_+> UAS-perRNAi.zip

https://github.com/srisaug/flywork/blob/main/RawData_Rivaetal_eLife2025_Fig4_M 122Bsplit-Gal4>+.zip

https://github.com/srisaug/flywork/blob/main/RawData_Rivaetal_eLife2025_Fig4_MB122Bsplit-Gal4>UAS-perRNAi.zip

https://github.com/srisaug/flywork/blob/main/RawData_Rivaetal_eLife2025_Figures1

Changes from the earlier version to the revised version

Abstract

An endogenous circadian clock controls many of the behavioral traits of Drosophila melanogaster. This “clock” relies on the activity of interconnected clusters of neurons that harbor the clock machinery. The hierarchy among clusters involved in the control of rest-activity cycles has been extensively studied. Sexually dimorphic behaviors, on the other hand, have received less attention. Even though egg-laying, a female characteristic behavior, has been shown to be rhythmic, it remains largely unexplored possibly due to metholodological constraints. The current study provides the first steps towards determining the neural substrates underlying the circadian control of egg-laying. We show that, whereas the lateral ventral neurons (LNvs) and the dorsal neurons (DNs) are dispensable, the lateral dorsal neurons (LNds) are necessary for rhythmic egg-laying. Systematically probing the Drosophila connectome for contacts between circadian clusters and oviposition-related neurons, we found no evidence of direct connections between LNvs or DNs and neurons recruited during oviposition. Conversely, we did find bidirectional connections between Cryptochrome (Cry) expressing LNd (Cry+ LNds) and oviposition related neurons. Taken together, these results reveal that Cry+ LNd neurons have a leading role in the control of the egg-laying rhythm in Drosophila females.

Introduction

Most organisms are capable of coordinating their physiology and behavior with the 24 hours of day/night cycling generated by the Earth’s rotation. These biological rhythms are driven by molecular clocks that are conserved across animals (Allada et al., 2001). In Drosophila melanogaster the core clock components comprise a transcription/translation feedback loop with four core proteins: CLOCK (CLK) CYCLE (CYC), PERIOD (PER), and TIMELESS (TIM) (Hardin, 2005; Tataroglu and Emery, 2015). In brief, CLK and CYC activate transcription of period and timeless genes, which, once translated into PER and TIM proteins, dimerize and translocate into the nucleus where they bind to CLK and CYC, thereby inhibiting their own transcription. This molecular feedback loop has a period of approximately 24 hours. In the Drosophila brain this molecular circadian clock is expressed in ∼ 242 neurons, which are organized in different clusters based in gene expression, anatomy and localization (Helfrich-Förster, 2005a; Reinhard et al., 2024; Yoshii et al., 2012). These clusters are: ventrolateral neurons (LNv; encompassing the small and large LNv groups), dorsolateral neurons (LNd), lateral posterior (LPN), and dorsal neurons [DN; separated in DN1, 2, and 3 and further subdivided into anterior (DN1a) and posterior (DN1p) DN1]. Rhythmic locomotor behavior depends primarily on the activity of the sLNvs, LNds, and DN1s, but all clusters contribute to different extents (Helfrich-Förster, 2005b). Within the lateral neuron group, the sLNv are important because they drive locomotor rhythmicity under free-running (DD) conditions (Grima et al., 2004; Stoleru et al., 2004), through the release of Pigment Dispersing Factor (PDF), a neuropeptide relevant for communication between clock neurons (Renn et al., 1999; Yoshii et al., 2009).

Egg-laying is one of the most important female behaviors since it has a profound impact on the fitness of a species. Egg laying is largely governed by successful mating, but is also influenced by circadian and seasonal rhythms and by environmental factors such as temperature and food availability, among others (Cury et al., 2019). The circadian rhythm of oviposition is one of the less studied rhythms in Drosophila, possibly due to the challenges involved in monitoring and recording this behavior. The periodic deposition of eggs involves a series of events ranging from the production of oocytes to the choice of the most appropriate substrate for the eggs (Allemand, 1976; Yang et al., 2008). The circadian nature of this behavior was revealed by its persistence under DD with a period around 24 h and a peak near night onset (McCabe and Birley, 1998a; Sheeba et al., 2001a). Egg-laying rhythmicity is temperature and nutrition compensated. (Howlader et al., 2006a). Moreover, oviposition is rhythmic in virgin females as well as in mated ones suggesting that this rhythm is not regulated by the act of mating and, instead, is endogenously driven (Menon et al., 2014a).

Although oviposition exhibits a circadian component, the molecular and neural substrates that govern this rhythm have only been partially described (Howlader et al., 2006b). Here, using an automated egg collection device developed in our lab, we examined the contribution of the molecular clock in specific neuronal clusters. By downregulating per in subsets of circadian neurons, we determined that even though impairing the molecular clock in the entire clock network reduced the power of the circadian rhythm of oviposition, restricting the disruption to DN1ps or LNvs did not (confirming results obtained by Howlader et al., 2006 for the sLNv (Howlader et al., 2006b)). Interestingly, egg-laying rhythms disappeared upon targeting a subset of LNd clocks (Cry+ LNd) through RNAi-mediated per downregulation, a condition where rhythmicity of locomotor activity patterns remained unaltered, suggesting a leading and very specific role of LNd clocks/neurons in the control of oviposition. Finally, the assessment of the synaptic connectivity between clock and oviposition controlling neurons using the Janelia hemibrain dataset revealed direct synaptic connections between this subset of LNds and oviposition neurons, which is consistent with the essential role of LNds neurons in the control of this behavior.

Results

Egg-laying is rhythmic when registered with a semiautomated egg collection device

Earlier version: Oviposition in Drosophila is one of the less studied behaviors regulated by the circadian clock. This is probably due to the difficulties involved in monitoring and recording it, and the current lack of standard devices to accomplish this. Egg collection and counting is usually done manually, making the experiments particularly demanding and labor-intensive. In this approach, eggs are typically collected every 4 hours (sometimes also every 2 hours), which usually implies transferring the fly to a new vial or extracting the food with the eggs and replacing it with fresh food in the same vial (Menon et al., 2014b). Either way this implies disturbing the fly several times a day which could alter the normal rhythm of oviposition, and demands the intervention of an experimenter every 4 hours during several days. In order to avoid this, we developed a semiautomated egg collection device where 21 flies are individually housed, and each enclosure is slowly shifted every 4 hours from one food patch to a new food patch. Once per day all food patches are collected and eggs are manually counted (see Material and Methods and Figure supplement 1 for more details).

Revised version: Oviposition in Drosophila is one of the less studied behaviors regulated by the circadian clock. This is probably due to the difficulties involved in monitoring and recording it, and the current lack of standard devices to accomplish this. Egg collection and counting is usually done manually, making the experiments particularly demanding and labor-intensive. In this approach, eggs are typically collected every 4 hours (sometimes also every 2 hours), which usually implies transferring the fly to a new vial or extracting the food with the eggs and replacing it with fresh food in the same vial (Menon et al., 2014b). Either way this implies disturbing the fly several times a day which could alter the normal rhythm of oviposition, and demands the intervention of an experimenter every 4 hours during several days. In order to avoid this, we developed a semiautomated egg collection device where 21 flies are individually housed, and each enclosure is slowly shifted every 4 hours from one food patch to a new food patch. Once per day all food patches are collected and eggs are manually counted (see Material and Methods; Figure supplement 1 and supplementary video for more details).

The egg records of fruit flies display several distinct features that complicate the assessment of rhythmicity. The recorded variable is discrete and typically takes only a few values, depending on the sampling frequency. The more often it is sampled, the fewer number of values the variable takes. Thus, it is only possible to sample it a few times per cycle (usually less than ten). For noisy signals, this feature makes the assessment of rhythmicity much more complicated, compared with signals that have similar noise, but are sampled more frequently (as is the case for activity records). The particularities of egg-laying in Drosophila add further complications. First, the number of eggs laid shows a clear decreasing trend over time. Second, this behavior appears to be particularly “noisy”, in the sense that the decision to lay an egg depends not only on the time of the day, but also on many other stimuli (Cury et al., 2019). As a consequence, any method that tries to assess the presence of a rhythm must be able to subtract, at least partially, the effects of trend and noise.

Earlier version: In order to achieve this, we have considered the register of the eggs laid by a “rhythmic” female as a periodic signal corrupted by random noise (whereas for arrhythmic females we assume the record is completely random). It is well known that light is a powerful zeitgeber for activity, synchronizing all the flies in a population to the light-dark (LD) cycle, and that even after lights off, activity remains synchronized between flies for quite a long time. Since this is caused by the synchronization of cellular clocks, we assume that a similar effect occurs for oviposition. Thus, under LD conditions, or in the first days of DD, it can be assumed that the periodic signal is the same for all flies, whereas the noisy component is different for each individual. In this case it is well known (see e.g. (van Drongelen, 2007)) that averaging of several “replicates” (individual flies, in our case) leads to an improvement of the signal-to-noise ratio allowing the emergence of the underlying rhythm. As this approach seems to be better suited for detecting noisy rhythms, in the rest of the paper the assessment of rhythmicity will be based in the record of eggs laid during a given period of time averaged over all females of each genotype. Figure 1 shows that, when eggs are collected with our device in LD, Canton-S females have a peak of egg deposition at the end of the day, as previously described in the literature (Bailly et al., 2023; Howlader et al., 2006a; Sheeba et al., 2001b) (Figure 1A). Figure 1E shows that only about 70% of individuals are rhythmic (with a period close to 24 h Figure 1F) although the average egg-laying profile is robust with a period around 24 h (Figure 1B). Under DD conditions, the percentages of rhythmic individuals are the same as in LD (Figure 1E) and the average profile also shows a robust rhythm with a period of 24 h (Figure 1D). It is well known that in locomotor activity experiments more than 90% of the flies are rhythmic in DD, and virtually all flies are rhythmic in LD (see e.g. Shindey et al. 2016, and Deppisch et al. 2022 (Deppisch et al., 2022; Shindey et al., 2017)). This shows that, at the individual level, egg-laying is much less robust than locomotor activity, compounding the difficulty of observing the influence of the circadian clock on this behavior.

Revised version: In order to achieve this, we have considered the register of the eggs laid by a “rhythmic” female as a periodic signal corrupted by random noise (whereas for arrhythmic females we assume the record is completely random). It is well known that light is a powerful zeitgeber for activity, synchronizing all the flies in a population to the light-dark (LD) cycle, and that even after lights off, activity remains synchronized between flies for quite a long time. Since this is caused by the synchronization of cellular clocks, we assume that a similar effect occurs for oviposition.. Thus, under LD conditions, or in the first days of DD, it can be assumed that the periodic signal is the same for all flies, whereas the noisy component is different for each individual. In this case it is well known (see e.g. (van Drongelen, 2007)) that averaging of several “replicates” (individual flies, in our case) leads to an improvement of the signal-to-noise ratio allowing the emergence of the underlying rhythm. As this approach seems to be better suited for detecting noisy rhythms, in the rest of the paper the assessment of rhythmicity will be based in the record of eggs laid during a given period of time averaged over all females of each genotype. Figure 1 shows that, when eggs are collected with our device in LD, Canton-S females have a peak of egg deposition at the end of the day, as previously described in the literature (Bailly et al., 2023; Howlader et al., 2006a; Sheeba et al., 2001b) (Figure 1A). Figure 1E shows that only about 70% of individuals are rhythmic (with a period close to 24 h Figure 1F) although the average egg-laying profile is robust with a period around 24 h (Figure 1B). Under DD conditions, the percentages of rhythmic individuals is the same as in LD (Figure 1E) and the average profile also shows a robust rhythm with a period of 24 h (Figure 1D). It is well known that in locomotor activity experiments more than 90% of the flies are rhythmic in DD, and virtually all flies are rhythmic in LD (see e.g. Shindey et al. 2016, and Deppisch et al. 2022 (Deppisch et al., 2022; Shindey et al., 2017)). This shows that, at the individual level, egg-laying is much less robust than locomotor activity, compounding the difficulty of observing the influence of the circadian clock on this behavior.

Oviposition in Drosophila is rhythmic when registered with our semiautomated egg collection device.

A, C, G, I: Average eggs collected as a function of time in four different experiment (Canton-S, LD: A, Canton-S, DD:C; YW: G; perS: I). The average was made with all flies (Strongly rhythmic, weakly rhythmic and arrhythmic). White and dark grey bars represent periods of lights on and off respectively (LD), whereas light gray bars represent subjective days, DD, (i.e times where lights were on at rearing, but are now off). B, D, H, J: Lomb-Scargle periodograms of all genotypes, made with all flies, strongly rhythmic, weakly rhythmic and arrhythmic. Red and green horizontal lines represent significances of 0.05 and 0.01, respectively. E, K: Percentage of females with rhythmic oviposition (E: Canton-S in LD and DD, K: YW and perS). The rhythmic flies include both strongly and weakly rhythmic flies. F, L: Period of oviposition rhythms for rhythmic individual flies (F: Canton-S in LD and DD, L: YW and perS). ns: non significant, ***p<0.001 (chi-squared test for scatter plots and proportions).

In order to check that our egg collection device does not in any way entrain egg laying in flies, we tested it with perS mutants (Konopka and Benzer, 1971). These flies are known to display a short period of locomotor activity, both in males (Konopka and Benzer, 1971) and in females (McCabe and Birley, 1998b). Additionally, it has been shown that oviposition in perS flies is rhythmic, with a period of 20.6+/- 0.9h (McCabe and Birley, 1998b). Our results for this mutant are in agreement with this (Figure 1J shows a period of the average at 18.8+/- 1.1 h), thus confirming that flies are not entrained by our egg collection device. It is important to notice that in the yellow white (YW) control strain, the individual rhythmicity was rather low (around 50% Figure 1K), but the average is strongly rhythmic (Figure 1H) thus confirming the need for an analysis method that can take into account the low signal-to-noise-ratio of this behavior.

Downregulation of per in clock neurons causes a dramatic reduction in the power of egg-laying rhythms

The circadian system of Drosophila comprises a central clock located in the brain that controls locomotor behavior, as well as peripheral clocks located in many tissues, that regulate distinct behaviors (Franco et al., 2018; Ito and Tomioka, 2016). For example, olfaction rhythms are controlled by peripheral clocks in the antennae, and eclosion rhythms are partially controlled by peripheral clocks in the prothoracic gland (Myers et al., 2003; Tanoue et al., 2004). Earlier work proposed that oviposition might also be controlled by a peripheral clock (Manjunatha et al., 2008). In order to test for this possibility, we used RNA interference (RNAi) to knock down the clock protein PER (Herrero et al., 2017; Zhang et al., 2021) in all clock neurons in the brain (using the Clk856 driver (Gummadova et al., 2009)), thereby disrupting their molecular clock. First, we tested the effectiveness of this tool by monitoring the locomotor activity of males with their endogenous clock silenced in all clock neurons (Clk856-Gal4>UAS-perRNAi). Supplementary Table 1 shows that these flies exhibited a significant decrease in locomotor rhythmicity compared to their genetic controls. Having confirmed the effectiveness of the RNA interference, we then monitored egg-laying in females from the previously mentioned lines.

As shown in Figure 2D,E, when PER is knocked down in all clock neurons oviposition rhythmicity is drastically reduced compared to controls Figure 2B,C. Notice, however, that the peak observed in Fig. 2C has not disappeared completely in 2E, and is in fact not far from the significance threshold. This could be due to some remaining rhythmicity in some individuals (although noise precludes the observation of this at the individual level). Interestingly, when locomotor activity is studied in male of this genotype, there is a fraction (less than 40%) of individuals that remain rhythmic in DD (see Supplementary Table 1).

Circadian rhythmicity of oviposition is dramatically reduced when the molecular clock is disrupted in all clock neurons, but not when only LNv or DN1 neurons are affected.

A, F, K, P: Schematic diagram of the neurons (painted in red) where the molecular clock has been disrupted. B, D, G, I, L, N, Q, S: Average eggs collected as a function of time. The average was made with all flies (strongly rhythmic, weakly rhythmic and arrhythmic). White and dark grey bars represent periods of lights on and off respectively (LD), whereas light gray bars represent subjective days, DD, (i.e times where lights were on at rearing, but are now off).C, E, H, J, M, O, R, T: Lomb-Scargle periodograms of all genotypes (made with all flies, rhythmic, weakly rhythmic and arrhythmic). Red and green horizontal lines represent significances of 0.05 and 0.01, respectively. B, C: +>UAS-perRNAi, n=22. D, E: Clk856-Gal4>UAS-perRNAi, n=26. G, H: +>UAS-perRNAi, n=18. I, J: PdfDicer-Gal4>UAS-perRNAi, n=34. L, M:+>UAS-perRNAi, n=38. N, O: Clk4.1-Gal4>UAS-perRNAi, n=40. Q, R: +>UAS-kir2.1, n=36. S, T:Clk4.1-Gal4>UAS-kir2.1, n=35.

This could be due either to the contribution of some peripheral oscillator/s relevant for the control of this behavior, or the presence of some residual PER in clock cells (Herrero et al., 2017).

LNv and DN1 neurons are not necessary for egg-laying rhythmicity

Earlier version: We next sought to establish the role of the most important circadian groups in the control of oviposition. We concentrated on three groups: LNv, DN1p and LNd neurons. The LNvs are known to be essential for the maintenance of circadian locomotor activity rhythms, and are the only ones that express the PDF neuropeptide (Renn et al., 1999; Yoshii et al., 2009). However, we found that perRNAi mediated disruption of the molecular clock in PDF+ cells does not abolish the time-of-day dependent oviposition (see Figure 2G,H,I,J). This is in line with previous results describing rhythmic oviposition after ablation of all PDF+ neurons (Howlader et al., 2006a). In addition, we observed a shortening of the egg-laying period when compared to controls, (20.74 h +/- 1.21 h vs 24.35 h+/-1.83 h respectively, Figure 2H,J). This shortening can also be observed at the individual level (Supplementary Figure 2S1). Taken together, these results suggest that these neurons may have some influence on the control of this behavior.

Revised version: We next sought to establish the role of the most important circadian groups in the control of oviposition. We concentrated on three groups: LNv, DN1p and LNd neurons. The LNvs are known to be essential for the maintenance of circadian locomotor activity rhythms, and are the only ones that express the PDF neuropeptide (Renn et al., 1999; Yoshii et al., 2009). However, we found that perRNAimediated disruption of the molecular clock in PDF+ cells does not abolish the time-of-day dependent oviposition (see Figure 2G,H,I,J). This is in line with previous results describing rhythmic oviposition after ablation of all PDF+ neurons (Howlader et al., 2006a). In addition, we observed a shortening of the egg-laying period when compared to controls, (20.74 h +/- 1.21 h vs 24.35 h+/-1.83 h respectively, Figure 2H,J). This shortening can also be observed at the individual level (Supplementary Figure 2S1). Taken together, these results suggest that these neurons may have some influence on the control of this behavior.

The DN1 neurons are not essential for the maintenance of locomotor rhythms in DD, but they contribute to the siesta (Guo et al., 2016) and generation of evening activity (Helfrich-Förster et al., 2007; Zhang et al., 2010) when the peak of egg deposition occurs. Interestingly, it has recently been shown that a posterior subset (DN1p) drives the circadian rhythm of oogenesis through the neuropeptide allatostatin C (Zhang et al., 2021). To examine whether the molecular clock of DN1p neurons is involved in the control of rhythmic oviposition, we evaluated the impact of PER downregulation on mated females expressing the Clk4.1 driver (which is expressed in ∼10 DN1ps per hemisphere (Zhang et al., 2010)). As Figure 2L,M,N,O shows, the disruption of the clock in these neurons does not alter the rhythmicity of egg-laying in DD. To further confirm these results, and in the process examine if some other features of the DN1p could contribute to the control of oviposition, we silenced these neurons through the expression of the inward rectifier potassium channel kir2.1 (Baines et al., 2001). Silencing DN1p neurons did not affect the rhythmicity of egg-laying in females which showed similar patterns to those exhibited by controls females (Figures 2Q,R,S,T). Taken together, these results show that neither LNv nor DN1p neurons play an important role for rhythmic oviposition.

The molecular clock in Cry+ LNd neurons is necessary for rhythmic egg-laying

Earlier version: The LNd neurons include 6 cells whose expression pattern is very heterogeneous (Hermann-Luibl and Helfrich-Förster, 2015; Ma et al., 2021). This group contains a subset of 3 Cry+ LNd neurons that express both cry (which encodes the light sensing protein cryptochrome) and pdfr (the PDF receptor), and 3 Cry- LNd neurons, that express neither of these genes (Im et al., 2011; Yoshii et al., 2008). Even though the LNds are not essential for rhythmic locomotor activity under constant conditions the Cry+ LNds, together with the 5-th sLNv, termed the “E oscillator”, are the main drivers of evening activity (Stoleru et al., 2004). These two LNd groups greatly differ in their connectivity to the rest of the brain, since Cry+ LNds have both many more outputs and inputs than the Cry- LNds (Shafer et al., 2022). In order to assess the role of the Cry+ LNds in the control of oviposition, we employed a Gal4 driver, Mai179-Gal4; pdf-Gal80 (Grima et al., 2004; Picot et al., 2007), to direct expression to E cells. This driver is expressed in the 5th sLNv, 3 Cry+ LNds, and also drives weak expression in a small subset of DN1s. As Figure 3 G shows, impairing clock function in Cry+ LNds results in a drastic reduction of rhythmic egg-laying compared with the genetic controls (Figure 3 C, E). To further confirm this observation, we used another approach to disrupt the molecular clock in those neurons, resorting to the very specific MB122B driver (Duhart et al., 2020; Guo et al., 2017). We also used a different UAS sequence, a dominant negative version of the clock gen cycle (Balaji Krishnan, 1999). As expected, under these conditions rhythmic egg laying was also severely compromised (Figure 3M, N), as compared with the genetic controls (Figure 3I, J, K, L). Moreover, when we used RNAi to knock down the clock protein PER with the MB122B driver, the results were very similar to those obtained with the Mai179-Gal4; pdf-Gal80 driver (Supplementary figure 3S1). Thus, we conclude that the molecular clock in Cry+ LNd neurons is necessary for the generation of egg-laying rhythms.

Revised version: The LNd neurons include 6 cells whose expression pattern is very heterogeneous (Hermann-Luibl and Helfrich-Förster, 2015; Ma et al., 2021). This group contains a subset of 3 Cry+ LNd neurons that express both cry (which encodes the light sensing protein cryptochrome) and pdfr (the PDF receptor), and 3 Cry- LNd neurons, that express neither of these genes (Im et al., 2011; Yoshii et al., 2008). Even though the LNds are not essential for rhythmic locomotor activity under constant conditions the Cry+ LNds, together with the 5-th sLNv, termed the “E oscillator”, are the main drivers of evening activity (Stoleru et al., 2004). These two LNd groups greatly differ in their connectivity to the rest of the brain, since Cry+ LNds have both many more outputs and inputs than the Cry- LNds (Shafer et al., 2022). In order to assess the role of the Cry+ LNds in the control of oviposition, we employed a Gal4 driver, Mai179-Gal4; pdf-Gal80 (Grima et al., 2004; Picot et al., 2007), to direct expression to E cells. This driver is expressed in the 5th sLNv, 3 Cry+ LNds, and also drives weak expression in a small subset of DN1s. As Figure 3 G shows, impairing clock function in Cry+ LNds results in a drastic reduction of rhythmic egg-laying compared with the genetic controls (Figure 3 C, E). To further confirm this observation, we used another approach to disrupt the molecular clock in those neurons, resorting to the very specific MB122B driver which is expressed in three LNds and 5th sLNv neurons (Duhart et al., 2020; Guo et al., 2017). We also used a different UAS sequence, a dominant negative version of the clock gene cycle (Balaji Krishnan, 1999). As expected, under these conditions rhythmic egg laying was also severely compromised (Figure 3M,N), as compared with the genetic controls (Figure 3I,J,K,L). Moreover, when we used RNAi to knock down the clock protein PER with the MB122B driver, the results were very similar to those obtained with the Mai179-Gal4; pdf-Gal80 driver (Supplementary Figure 3S1). Thus, we conclude that the molecular clock in Cry+ LNd neurons is necessary for the generation of egg-laying rhythms.

Disruption of the molecular clock in Cry+ LNd neurons drastically reduces the circadian rhythmicity of oviposition.

A, H: Schematic diagram of the neurons (painted in red) where the molecular clocks have been disrupted. B, D, F, I, K, M: Average eggs collected as a function of time. The average was made with all flies (Strongly rhythmic, weakly rhythmic and arrhythmic). White and dark grey bars represent periods of lights on and off respectively (LD), whereas light gray bars represent subjective days, DD, (i.e times where lights were on at rearing, but are now off).C, E, G, J, L, N: Lomb-Scargle periodograms of all genotypes (made with all flies, rhythmic, weakly rhythmic and arrhythmic). Red and green horizontal lines represent significances of 0.05 and 0.01, respectively. B, C: +>UAS-perRNAi, n= 15. D, E: Mai179-Gal4;pdf-Gal80>+, n= 17. F, G: Mai179-Gal4;pdf-Gal80>UAS-perRNAi, n=28. I, J: +>UAS-cycDN, n=15. K, L: MB122Bsplit-Gal4>+, n=14. M, N: MB122Bsplit-Gal4>UAS-cycDN, n=30.

Earlier version: The appearance of two non-significant peaks in Figure 3 G (and Supp. Fig 3S1) might suggest that these genetic manipulations of the LNds neurons generate weakened complex rhythms (Beckwith and Ceriani, 2015) instead of complete arrhythmicity. This would mean that, at the individual level, there should be either many flies with a corresponding complex rhythm, or two weakly rhythmic subpopulations, each one accounting for one of the peaks of the average periodogram. However, an examination of the individual periodograms (Supplementary Figures 3S2 and 3S3) does not provide convincing evidence to support any of these two hypotheses. On the other hand, when the method is applied to synthetic, completely random time series, it can generate periodograms with two non significant peaks (as shown in panels 2,3, and 4 of Figure S12 of the Supplementary Material). Thus, we favor the hypothesis that the two peaks present in Figure 3 G (and Supplementary Figure 3S1) are spurious.

Revised version: The appearance of two non-significant peaks in Figure 3 G (and Supp. Fig 3S1) might suggest that these genetic manipulations of the LNds neurons generate weakened complex rhythms (Beckwith and Ceriani, 2015) instead of complete arrhythmicity. This would mean that, at the individual level, there should be either many flies with a corresponding complex rhythm, or two weakly rhythmic subpopulations, each one accounting for one of the peaks of the average periodogram. However, an examination of the individual periodograms (Supplementary Figures 3S2 to 3S7) does not provide convincing evidence to support any of these two hypotheses. On the other hand, when the method is applied to synthetic, completely random time series, it can generate periodograms with two non significant peaks (as shown in panels 2, 3, and 4 of Figure S12 of the Supplementary Material). Thus, we favor the hypothesis that the two peaks present in Figure 3 G (and Supplementary Figure 3S1) are spurious.

It has recently been shown that disrupting the molecular clock in E cells by means of the MB122B driver did not alter locomotor activity rhythms (Bulthuis et al., 2019). This would imply that locomotor activity and oviposition rhythms are controlled by distinct groups of clock neurons. However, Bulthuis et al. (2019) only assessed the locomotor activity of male flies, and it is well known that, unlike males, females undergo important changes in their activity after mating, such as the loss of the midday siesta (Elwyn Isaac et al., 2010) or the morning anticipation (Riva et al., 2022). For these reasons, we decided to examine the influence of the molecular clock of E cells directly on the locomotor activity of mated females. The standard setup for assessing locomotor activity is the Drosophila Activity Monitoring (DAM) system (Pfeiffenberger et al., 2010), where flies are housed in small glass tubes (of 5mm in diameter). Such arena does not offer an adequate environment for testing flies that lay eggs during several days; instead, we monitored activity using an alternative system developed in our lab, where flies are housed in much larger cubicles and their activity is recorded by a video system (see Methods and (Riva et al., 2022)). Interestingly, mated females with disrupted E cell clocks were as rhythmic as control flies (Figure 4), showing that there is no sex-specific contribution of this cluster to the control of rhythmic rest-activity cycles, which opens the possibility that E cells differently contribute to the control of oviposition and locomotor activity. In other words, our results show that these two important rhythmic behaviors are primarily controlled by distinct groups of clock neurons.

Disruption of the molecular clock in Cry+ LNd neurons does not alter the circadian rhythmicity of locomotor activity of mated female flies.

A: Average activity recording during 3 days in LD and 7 in DD. Light grey bars represent subjective days. B: Percent of rhythmic flies in DD. C: Periods of locomotor activity in DD of individually rhythmic flies. Mean ± SEM indicated. Each dot represents one fly. ns: non significant, *: p<0.05 (chi-squared test for the comparison of proportions, and one-way-ANOVA with post-hoc Tukey for the scatter plots.

Cry+ LNd neurons directly contact oviIN

If the circadian clock influences oviposition, at least some clock neurons should be connected, directly or indirectly, to the neurons in the brain that control the motor program leading to egg deposition. These neurons, that were recently characterized (Vijayan et al., 2023a; Wang et al., 2020; Zhou et al., 2014), include: the oviposition descending neurons (oviDN, 5 per hemisphere), the oviposition excitatory neurons (oviEN, 1 per hemisphere), the oviposition inhibitory neurons (oviIN, 1 per hemisphere), and pC1 neurons (5 per hemisphere). The best tool in hand to systematically probe the connectivity between circadian clock and oviposition-related neurons is the Drosophila hemibrain connectome (Scheffer et al., 2020). This is the result of the FlyEM connectomics project, which used electron microscopy images of the brain of a 5-day old female fly to reconstruct the connectivity map (or connectome) of ∼25000 neurons, mostly from the central part and the right lateral part of the brain.

28 of these neurons have been identified as clock neurons, and are located in the right part of the brain (except for one LPN (Scheffer et al., 2020)). These include the best characterized circadian groups (sLNvs, lLNvs, LNds). For other groups, such as the DN1s and LPNs, only a minor fraction have been identified, whereas for groups DN2 and DN3 no neurons have yet been identified in this volume.

Interestingly, almost all oviposition-related neurons have been identified in the hemibrain (only two oviDN on the left hemisphere are still missing), and very recently a couple of additional groups called U and G, comprising 2 and 5 neurons each, respectively, have also been shown to be important for oviposition and mapped to the connectome (Vijayan et al., 2023b). This makes a total of 25 oviposition-related neurons identified thus far.

To address the possibility of direct connections between clock neurons and neurons involved in egg laying, we probed the hemibrain connectome for these connections. The full results are displayed in Supplementary Tables 2 and 3, but in short, we found that whereas there are no connections between LNv or DN1 neurons and oviposition neurons, LNds contact oviIN and pC1 neurons (Figure 5), further supporting the experimental findings described in previous sections. Interestingly, different subgroups of E cells (Shafer et al., 2022) show distinct patterns of connectivity to oviposition neurons: E3 neurons (comprising Cry- LNd1, LNd2 and LNd3) only contact the pC1 group, while E2 neurons (comprising Cry+ LNd4 and LNd5) only contact the oviIN neurons (Figure 5). On the other hand, the E1 neurons (comprising Cry+ LNd6 and 5-th sLNv) make no contacts with oviposition related neurons. This evidence suggests that the subset of E cells responsible for the loss of rhythmic oviposition are likely E2 neurons.

Direct synaptic connections between circadian clock neurons and oviposition-related neurons in the hemibrain dataset.

A: Schematic diagrams showing the different neuron clusters analyzed (in color). B: Connection between Cry+ LNd and oviIN neurons. C: Connection between Cry-LNd 2and pC1b neurons. D: Network representation of the connectivity between Cry+ LNd and oviIN neurons. E: Network representation of the connectivity between Cry- LNd and pC1b neurons. Numbers give the number of synaptic contacts, which represent the strength of the connections. Only intermediate (between 3 and 9 synaptic contacts) and strong (>9 synaptic contacts) connections are considered.

Interestingly, we noticed that connections between clock and oviposition neurons are bidirectional, although the E2-oviIN connection is stronger in the direction from the clock to oviposition neurons. This analysis suggests that, even though some feedback is anticipated, the E2 neurons are likely the ones providing temporal information to the oviposition circuit.

Additionally, aside from the circadian groups considered so far, we found strong, bidirectional connections between the clock neurons LPN (Reinhard et al., 2022) and the oviIN, with the strongest connections going from oviIN to LPN (Supplementary Tables 2 and 3 and Supplementary Figure 5S1).

Discussion

The oviposition rhythm is one of the less studied behavioral traits of Drosophila. One possible reason is merely technical, since egg collection and counting are a very laborious processes and there are no standard devices (just as the DAM system for locomotor activity) to reproducibly evaluate it. Data analysis and rhythmicity assessments are also particularly difficult because the presence or absence of a rhythm must be established with only approximately 6 data points per cycle (in contrast to locomotor activity time series, where hundreds of data points per cycle are obtained), which can take only a few discrete values (typically, from 0 to 10). The analysis is further complicated by the fact that the daily average of eggs laid decreases with time (Kaufman and Demerec, 1942). To tackle these limitations, we developed a semiautomated method for egg collection and a new pipeline for data analysis. We observed that the oviposition rhythm is less consolidated than locomotor activity rhythms. Surprisingly, this happens even in LD, where time of day cues are provided by light, the most important zeitgeber for diurnal animals. Egg-laying is even less rhythmic in DD since the information from the endogenous clock is one among many internal and external factors that influence the decision of laying an egg. In a sense, the oviposition record of a single fly can be considered as periodic signal with strong noise added to it. Such data might not even look rhythmic when analyzed individually. However, if we assume that the periodic component is the same for all flies in a population, and the noise is different for each individual, the periodic component can be extracted by averaging the oviposition data. This led us to prefer periodograms of averaged data for the assessment of the rhythmicity of a given genotype.

Even though there have been some advances in the understanding of the relationship between circadian clock and oviposition (Howlader and Sharma, 2006; Manjunatha et al., 2008), there is still no information about the relative importance of the different neuronal groups in driving this behavior. Here, we downregulated the molecular clock in subsets of circadian neurons in order to establish their role in the control of circadian oviposition. First, we found that animals with downregulated per expression in all clock neurons of the brain display a drastic reduction of rhythmic oviposition, underscoring a key role of the central clock (as compared with peripheral clocks) in the control of the timing of oviposition.

Even though we also disrupted the molecular clock in PDF-expressing neurons as well as in the DN1p group, egg-laying rhythms were not abolished, in agreement with previous reports (Howlader et al., 2006a). We cannot rule out a contribution of PDF-expressing neurons, since we observed a shortening of the oviposition period. Interestingly, a previous work (Howlader et al., 2006b) reported a similar period shortening in pdf⁰¹ mutant flies, but not when PDF-expressing neurons were ablated. Taken together, these results show that communication between either group and any other neurons is not necessary for rhythmic oviposition. This is in stark contrast with the situation for locomotor activity, which has been shown to become arrhythmic when PDF+ cells are ablated (Grima et al., 2004; Stoleru et al., 2004). Interestingly, however, it has been recently shown that disrupting the molecular clock in these same neurons does not alter locomotor rhythmicity (Delventhal et al., 2019; Schlichting et al., 2019) suggesting a more complex scenario.

Since it has been shown that the DN1p generate the rhythm of oogenesis (Zhang et al., 2021), one might wonder why the disruption of their molecular clock or their electrical silencing does not also abolish the rhythm of oviposition. It is important to understand that the suppression of the oogenesis rhythm does not necessarily have an impact on rhythmic egg laying. The persistence of oviposition rhythm suggests that it does not depend on the rhythmic production of oocytes, but is instead an intrinsic property of the motor program which is the hypothesis more supported by our results.

Somewhat surprisingly, the disruption of the molecular clock in Cry + LNd neurons led to an almost complete loss of the oviposition rhythm. In contrast, downregulating the clock in Cry+ LNds did not affect rhythmic locomotor activity in mated females, underscoring a degree of specificity. Interestingly, in males the Cry+ LNds have been singled out as responsible for driving the activity peak at dusk (evening peak, (Guo et al., 2017)), which is the time when, in females, egg laying has a peak. As far as we know, oviposition is the first female-specific behaviour specifically controlled by the LNds. In males the LNds have been shown to be involved in the circadian control of the expression of sex-specific fat body genes (Erion et al., 2016; Fujii et al., 2008).

The neural circuitry that controls the egg laying process has just begun to be unraveled (Vijayan et al., 2023a; Wang et al., 2020; Zhou et al., 2014). It includes the oviDNs, oviINs, oviENs, groups U and G, and pC1 neurons. Using the hemibrain connectome we found that neither the identified DN1s nor the LNv neurons directly contact any neurons ascribed to the oviposition circuit; the LNd cluster, on the other hand, exhibits synaptic contacts with pC1 and oviIN neurons, providing the anatomical substrate for our experimental results. Furthermore, there is a clear separation of roles inside the LNd cluster, since the oviINs are only contacted by two Cry+ LNds (the so-called E2 cells (Shafer et al., 2022)) and the pC1s are only contacted by the three Cry- LNds (E3 cells (Shafer et al., 2022)). There is also a functional separation within the pC1 cluster, since the only pC1 neurons connected E3 neurons are the pC1b neurons from both right and left hemispheres (Figure 5). The resulting clock/oviposition network is summarized in Figure 6.

Scheme of connections between circadian clock and oviposition-related neuron clusters of the same hemisphere in the hemibrain dataset.

Each circle represents a neuron cluster, comprising different numbers of neurons (some clusters comprise only one neuron). Clusters and connections involved in the connectivity between circadian clock and oviposition sets have been colored. In each connection the arrow points to the post-synaptic cluster. OviIN neurons are bidirectionally connected to every neuron of the clusters inside the square.

Our data posits the E2 LNds-oviIN connection as the candidate to mediate the timing of egg laying rhythms. The oviIN neurons are two of the most important neuronal hubs included in the current version of the Drosophila connectome, receiving strong inputs of hundreds of other neurons. Thus, the oviINs integrate the information from the circadian clock with the information coming from many other sources in order to signal to the oviposition circuit. Under such scenario the circadian clock would compete with many other sources to set the threshold for egg deposition (Vijayan et al., 2023a).

We hypothesize that the E3-pC1b connection is less important for conveying circadian information to the oviposition circuit because the pC1b are located upstream within the oviposition circuit that controls the motor program, being connected only to the oviIN neurons (Wang et al., 2020). On the other hand, since pC1 neurons control female receptivity, it is tempting to speculate that this behavior may have a circadian rhythm, that could be generated through the E3-pC1b connection. Interestingly, the female mating rate has been shown to display circadian rhythmicity (Sakai and Ishida, 2001), and very recently it has been proposed that pC1b neurons are involved in the control of female copulation rate (Li et al., 2023). Conversely, given that the E3-pC1b connection is bidirectional, it could also be instrumental in conveying information about mating status to the circadian clock, therefore producing the changes in temporal organization that females undergo after mating (Delbare et al., 2023; Elwyn Isaac et al., 2010; Riva et al., 2022).

The function of the LPN neurons in the circadian clock is still unclear (Reinhard et al., 2022) and therefore no clear role in rhythmic oviposition can be ascribed to the LPN-oviIN connection. On the other hand, since there are many more oviIN inputs to the LPN than outputs (135 vs 16 synaptic connections), it seems more likely that this connection conveys information to the circadian clock from those processes in which the oviINs are involved.

The information provided by the Drosophila connectome is only the first step towards the understanding of the influence of the circadian clock on oviposition. Future work will be necessary to test the functionality of these connections as well as the role of different neuropeptides (sNPF, ITP, PDF) and neurotransmitters (such as acetylcholine and glutamate) in the control of egg laying behavior.

Material and methods

Fly strains

Earlier version: All fly strains used in this study are detailed in Table 1. Flies were reared and maintained on standard cornmeal/agar medium at 25 °C and 60% humidity in a 12 hr:12 hr LD cycle unless stated otherwise. For oviposition and locomotor activity experiments we used a different medium, named banana medium, prepared with 200g of banana, 25g of barley, 36g of black sugar, 35 g of yeast, 12.5g of agar and 2g of Nipagin per liter of water 15.

Revised version: All fly strains used in this study are detailed in Table 1. Flies were reared and maintained on standard cornmeal/agar medium at 25 °C and 60% humidity in a 12 hr:12 hr LD cycle unless stated otherwise. For oviposition and locomotor activity experiments we used a different medium, named banana medium, prepared with 200g of banana, 25g of barley, 36g of black sugar, 35 g of yeast, 12.5g of agar and 2g of Nipagin per liter of water.

Semiautomated egg collection device

Earlier version: The egg collection device (Supplementary Figure 1) consists of a wooden basis with a mechanical arm fixed to it, and plastic sets of tracks with fly chambers that can be replaced at will. Each “set of tracks” is a 13.9 x 19.5 cm 3D-printed piece with 7 tracks. Inside each track there are 6 equidistant wells measuring 15 x 20 mm, which are filled with banana medium. Each fly chamber (17 x 27 mm) has a transparent roof, and no floor, so that the fly inside it is in contact with the food wells, when the chamber is placed on the track so that it can slide along it. 3 sets of tracks are placed on the wooden base, and the chambers are moved together from well to well by a mechanical arm, which is moved by a stepper motor controlled by an Arduino UNO. The chambers are displaced from well to well every 4 hours. More frequent sampling gives rise to less consistent rhythmic patterns. Survival in this egg collection device is in general high: on average, more than 80% of the flies are alive at the end of an egg collection experiment lasting one week.

Revised version: The egg collection device (Supplementary Figure 1 S1 and supplementary video) consists of a wooden basis with a mechanical arm fixed to it, and plastic sets of tracks with fly chambers that can be replaced at will. Each “set of tracks” is a 13.9 x 19.5 cm 3D-printed piece with 7 tracks. Inside each track there are 6 equidistant wells measuring 15 x 20 mm, which are filled with banana medium. Each fly chamber (17 x 27 mm) has a transparent roof, and no floor, so that the fly inside it is in contact with the food wells, when the chamber is placed on the track so that it can slide along it. 3 sets of tracks are placed on the wooden base, and the chambers are moved together from well to well by a mechanical arm, which is moved by a stepper motor controlled by an Arduino UNO. The chambers are displaced from well to well every 4 hours. More frequent sampling gives rise to less consistent rhythmic patterns. Survival in this egg collection device is in general high: on average, more than 80% of the flies are alive at the end of an egg collection experiment lasting one week.

Behavioral assays

The egg deposition behavior of females was analyzed individually at 25°C and 60% humidity. Before starting the assay, six 0-5 day old virgin females and five males of the desired genotype were anesthetized using CO2 and introduced into a vial with 10mL of standard food for 72 hours. This allowed for the crossing of 50-80 female flies with their respective males per experiment. Subsequently, the resulting gravid females from the crosses were placed individually in vials containing a plastic spoon with 1.5mL banana medium along with a drop of yeast for one day. The eggs deposited by each of the females are counted under a binocular stereoscopic magnifier (Lancet Instruments) and the 21 best egg layers are selected to conduct the circadian oviposition experiment. This prior selection of the best egg layers is performed because, of the total females put to mate with males, nearly half do not lay eggs or lay very few.

21 of these females were housed in the chambers of each apparatus, which was then introduced in an incubator at constant temperature (25 °C). The first 1 or 2 days the light regime was LD (12:12) in order for the flies to adapt to their new environment, and then they entered a DD regime for a given number of days. The sets of tracks were retired and replaced by sets with new food every 20, 24 or 28 hours (in order to avoid providing an entrainment signal to the flies). The eggs in each food well of each of the retired tracks were counted under a binocular stereoscopic magnifier (Lancet Instruments), and registered.

For the locomotor activity assays we used the setup used in previous works (Riva et al., 2022). Briefly, the flies are housed in translucid tracks and their movement is registered with video cameras above the tracks, connected to a computer, where the position of each fly is extracted from the video (with ad hoc programmed software) and registered. Locomotor activity of mated females was monitored for 3 days in LD conditions and then transferred to DD conditions for 7 days. Every 4 or 5 days the flies were transferred to tracks with fresh food, because after that time the appearance of larvae hindered the video tracking of the females.

Data analysis

Earlier version: After each oviposition experiment lasting N days the data consisted of a time series E of 6N points for each fly. A new time series was generated by averaging the individual series. Since the number of eggs laid by a mated female tend to show a downward trend, we proceeded as follows, in order to detrend the data (see the Supplementary Material for further details). First, a moving average of the data is performed, with a 6 point window, and a new time series T is obtained. In principle, T is a good approximation to the trend of the data. Then, a new, detrended, time series D is generated by pointwise dividing the two series (i.e. D(i)=E(i)/T(i), where i indexes the points of each series).

Revised version: After each oviposition experiment lasting N days the data consisted of a time series E of 6N points for each fly. A new time series was generated by averaging the individual series. Since the number of eggs laid by a mated female tend to show a downward trend, we proceeded as follows, in order to detrend the data (see the Supplementary Material for further details). First, a moving average of the data is performed, with a 6 point window, and a new time series T is obtained. In principle, T is a good approximation to the trend of the data. Then, a new, detrended, time series D is generated by pointwise dividing the two series (i.e. D(i)=E(i)/T(i), where i indexes the points of each series). In order to assess the rhythmicity of a genotype, we averaged the detrended time series and performed a Lomb-Scargle periodogram (Ruf, 1999; VanderPlas, 2018). For this we used package Lomb (version 2.1.0), from the R Statistical Software (v4.1.2; R Core Team 2021). The significance lines were calculated by repeatedly randomizing the time series (using function randslp in Lomb R package). A genotype was considered rhythmic if the periodogram had one peak between 16 and 32 hs, and it was above the p=0.05 significance line.

Removed in the revised version: In order to assess the rhythmicity of a genotype, we averaged the detrended time series and performed a Lomb-Scargle periodogram (Ruf, 1999; VanderPlas, 2018). For this we used package Lomb (version 2.1.0), from the R Statistical Software (v4.1.2; R Core Team 2021). The significance lines were calculated by repeatedly randomizing the time series (using function randslp in Lomb R package). A genotype was considered rhythmic if the periodogram had one peak between 16 and 32 hs, and it was above the p=0.05 significance line.

Earlier version: The graphs showing individual periods were made by obtaining the individual periods using the Lomb-Scargle periodogram described above for each particular individual. For the individual periods shown in the figures, we added the weakly rhythmic category, corresponding to flies whose periodograms displayed only one peak between 18 and 32 hs, and a power larger than 0.2.

Revised version: The graphs showing individual periods were made by obtaining the individual periods using the Lomb-Scargle periodogram described above for each particular individual. For the individual periods shown in the figures, we added the weakly rhythmic category, corresponding to flies whose periodograms displayed only one peak between 18 and 32 hs, and a power larger than 0.2. More details about the method used can be found in the Supplementary Material.

Removed in the revised version: More details about the method used can be found in the Supplementary Material 2.

For locomotor activity experiments, these files obtained from the video tracking were processed with an analysis software we developed (in Bash) which provides statistics for activity (position, distance traveled, etc.) (Riva et al., 2022).

Analysis of connectome data

Earlier version: The data used to determine the connectivity between the circadian clock and the oviposition circuits come from the Hemibrain dataset (version 1.2.1) made publicly available by Janelia Research Campus (Scheffer et al., 2020). To access the data we used the NeuPrintExplorer Web tool (https://neuprint.janelia.org/). The result of our queries is summarized in Supplementary Tables 2 and 3. The neuron IDs are as given in the hemibrain dataset, just as the instance names of the oviposition related neurons. For circadian clock neurons we used the names given by Shafer et al. (Shafer et al., 2022). We only replaced LPN-4 by LPN-L to stress the fact that this neuron is placed in the left hemisphere (all the others are in the right hemisphere). The strength of each connection was quantified using the criteria of Scheffer et al. (Scheffer et al., 2020). Connections with more than 9 synapses were considered strong, connections having between 3 and 9 synapses were considered as having intermediate strength, and connections with 0,1, or 2 synapses were considered as weak. Because weak connections are prone to error (Scheffer et al., 2020) we excluded them from the analyses. The figures that show connectivity between sets of neurons were drawn using Cytoscape version 3.10.1 (Shannon et al., 2003).

Revised version: The data used to determine the connectivity between the circadian clock and the oviposition circuits come from the Hemibrain dataset (version 1.2.1) made publicly available by Janelia Research Campus (Scheffer et al., 2020). To access the data we used the NeuPrintExplorer Web tool (https://neuprint.janelia.org/). The result of our queries is summarized in Supplementary Tables 2 and Table 3. The neuron IDs are as given in the hemibrain dataset, just as the instance names of the oviposition related neurons. For circadian clock neurons we used the names given by Shafer et al. (Shafer et al., 2022). We only replaced LPN-4 by LPN-L to stress the fact that this neuron is placed in the left hemisphere (all the others are in the right hemisphere). The strength of each connection was quantified using the criteria of Scheffer et al. (Scheffer et al., 2020). Connections with more than 9 synapses were considered strong, connections having between 3 and 9 synapses were considered as having intermediate strength, and connections with 0,1, or 2 synapses were considered as weak. Because weak connections are prone to error (Scheffer et al., 2020) we excluded them from the analyses. The figures that show connectivity between sets of neurons were drawn using Cytoscape version 3.10.1 (Shannon et al., 2003).

Statistical analysis

The following statistical analyses were used for the comparison between genotypes of individual periods (scatter plots). First, we tested data for normality with D’Agostino & Pearson test. Once the normality of the data was confirmed, we used a unpaired t-test for test differences between means for experiments where 2 groups were compared and a one-way ANOVA for experiments where 3 groups were compared. When data was not normally distributed, we applied a Mann-Whitney test for test differences between means. When significant differences were found, we conducted a post-hoc Tukey test with a correction for multiple comparisons. In all the analyses we used Two-Tailed p-values. To perform all the tests, we used GraphPad Prism (Boston, Massachusetts USA, www.graphpad.com) and R Statistical Software (v4.1.2; R Core Team 2021). p< 0.05 was considered statistically significant. In dot plots, horizontal lines represent the mean value; error bars depict the standard error of the mean. For the comparison between genotypes of the proportion of rhythmic individuals over the total (proportion plots), we used the ‘N-1’ Chi-squared test as recommended by Campbell (2007) and Richardson (2011). To perform these tests, we use MedCalc Software Ltd. Comparison of proportions calculator (https://www.med-calc.org/calc/comparison_of_proportions.php (Version 23.2.7; accessed June 19, 2025). No statistical methods were used to determine sample size. Sample sizes are similar to those generally used in this field of research. Samples were not randomized and analyzers were not blind to the experimental conditions. In all the figures we show results of two or three independent experiments.